<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab09/lab09_td.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 9: Regresión

¡Bienvenidos al Laboratorio 9!

Hoy practicaremos un poco con la regresión lineal. Puedes encontrar más información sobre este tema en
[Capítulo 15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-regression-line).

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

```python
# Ejecute esta celda, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.

import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

# 1. ¿Qué tan fiel es Old Faithful?

Old Faithful es un géiser en el Parque Nacional de Yellowstone que es famoso por hacer erupción en un horario bastante regular. Ejecute la celda a continuación para ver a Old Faithful en acción.

```python
# Para los curiosos: así es como mostrar un vídeo de YouTube en un
# Cuaderno Jupyter. El argumento de YouTubeVideo es la parte.
# de la URL (llamada "parámetro de consulta") que identifica el
# video. Por ejemplo, la URL completa de este vídeo es:

#   https://www.youtube.com/watch?v=wE8NDuzt8eg
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Algunas de las erupciones de Old Faithful duran más que otras. Siempre que hay una erupción prolongada, suele ir seguida de una espera aún más larga antes de la siguiente erupción. Si visita Yellowstone, es posible que desee predecir cuándo ocurrirá la próxima erupción, para poder ver el resto del parque en lugar de esperar junto al géiser.

Hoy utilizaremos un conjunto de datos sobre la duración de las erupciones y los tiempos de espera para ver si podemos hacer tales predicciones con precisión mediante regresión lineal.

El conjunto de datos tiene una fila para cada erupción observada. Incluye las siguientes columnas:
- `duración`: duración de la erupción, en minutos
- `espera`: Tiempo entre esta erupción y la siguiente, también en minutos

Ejecute la siguiente celda para cargar el conjunto de datos.

```python
faithful = Table.read_table("faithful.csv")
faithful
```

**Pregunta 1.0.** Las siguientes afirmaciones son los pasos desordenados de la regresión lineal.

1. Calcule los parámetros de la recta de regresión: la pendiente y la intersección.
2. Evalúe la recta de regresión calculando el RMSE de la recta y analizando la gráfica de residuos.
3. Utilice la línea de regresión para generar predicciones para cada valor de x.
4. Determine si la regresión lineal es un método razonable visualizando sus datos y calculando el coeficiente de correlación.

Haga una matriz llamada `least_squares_order` que contenga el orden correcto de un análisis de regresión lineal, donde el primer elemento de la matriz es el primer paso de un análisis de regresión lineal y el último elemento de la matriz es el último paso de un análisis de regresión lineal.

```python
least_squares_order = ...
```

```python
grader.check("q1_0")
```

Nos gustaría utilizar la regresión lineal para hacer predicciones, pero eso no funcionará bien si los datos no están relacionados de forma lineal. Para comprobarlo, deberíamos mirar los datos.

**Pregunta 1.1.** Haz un diagrama de dispersión de los datos. Es convencional poner la columna que queremos predecir en el eje vertical y la otra columna en el eje horizontal.


```python
...
```

**Pregunta 1.2.** ¿La duración de la erupción y el tiempo de espera están relacionados de manera aproximadamente lineal según el diagrama de dispersión anterior? ¿Es esta relación positiva?


_Escriba su respuesta aquí, reemplazando este texto._

Continuaremos asumiendo que están relacionados linealmente, por lo que es razonable utilizar la regresión lineal para analizar estos datos.

A continuación nos gustaría trazar los datos en unidades estándar. Si no recuerda la definición de unidades estándar, la sección  del libro de texto [14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) puede resultarle útil.

**Pregunta 1.3.** Calcule la media y la desviación estándar de la duración de las erupciones y los tiempos de espera. **Luego** cree una tabla llamada `faithful_standard` que contenga la duración de las erupciones y los tiempos de espera en unidades estándar. Las columnas deben denominarse "duración (unidades estándar)" y "espera (unidades estándar)".


```python
duration_mean = ...
duration_std = ...
wait_mean = ...
wait_std = ...

faithful_standard = Table().with_columns(
    "duration (standard units)", ...,
    "wait (standard units)", ...)
faithful_standard
```

```python

```

```python
grader.check("q1_3")
```

**Pregunta 1.4.** Traza los datos nuevamente, pero esta vez en unidades estándar.


```python
...
```

¡Notarás que esta trama se ve igual que la anterior! Sin embargo, los datos y los ejes tienen una escala diferente. Por eso es importante leer las marcas en los ejes.

**Pregunta 1.5.** Entre los siguientes números, ¿cuál cree que es el más cercano a la correlación entre la duración de la erupción y el tiempo de espera en este conjunto de datos?

1. -1
2. 0
3. 1

Asigne "correlación" al número correspondiente a su conjetura (ya sea 1, 2 o 3).


```python
correlation = ...
```

```python
grader.check("q1_5")
```

**Pregunta 1.6.** Calcule el coeficiente de correlación: `r`.

*Pista:* Utilice `faithful_standard`. La sección [15.1](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) explica cómo hacer esto.



```python
r = ...
r
```

```python
grader.check("q1_6")
```

## 2. La línea de regresión
Recuerde que la **correlación** es la **pendiente de la línea de regresión cuando los datos se expresan en unidades estándar**.

La siguiente celda traza la línea de regresión en unidades estándar:

$$\text{tiempo de espera en unidades estándar} = r \times \text{duración de la erupción en unidades estándar}$$

Luego, vuelve a trazar los datos en unidades estándar para compararlos.

```python
def plot_data_and_line(dataset, x, y, point_0, point_1):
    """Makes a scatter plot of the dataset, along with a line passing through two points."""
    dataset.scatter(x, y, label="data")
    xs, ys = zip(point_0, point_1)
    plots.plot(xs, ys, label="regression line")
    plots.legend(bbox_to_anchor=(1.5,.8))

plot_data_and_line(faithful_standard,
                   "duration (standard units)",
                   "wait (standard units)",
                   [-2, -2*r],
                   [2, 2*r])
```

¿Cómo tomarías un punto en unidades estándar y lo convertirías nuevamente a unidades originales? Tendríamos que "estirar" su posición horizontal por `duration_std` y su posición vertical por `wait_std`. Eso significa que le pasaría lo mismo a la pendiente de la línea.

Estirar una línea horizontalmente la hace menos empinada, por lo que dividimos la pendiente por el factor de estiramiento horizontal. Estirar una línea verticalmente la hace más empinada, por lo que multiplicamos la pendiente por el factor de estiramiento vertical.

**Pregunta 2.1.** Calcule la pendiente de la recta de regresión en unidades originales y asígnela a "pendiente".

(Si la explicación del "estiramiento" no es intuitiva, consulte la sección [15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) en el libro de texto.)


```python
slope = ...
slope
```

```python
grader.check("q2_1")
```

Sabemos que la recta de regresión pasa por el punto `(duración_media, espera_media)`. Recuerde que la ecuación de la recta de regresión en las unidades originales es:

$$ \text{tiempo de espera} = \text{pendiente} \cdot \text{duración de la erupción} + (- \text{pendiente} \cdot \text{duración media + espera media})$$


**Pregunta 2.2.** Calcule la intersección en unidades originales y asígnela a "intersección". [Sección 15.2.5](https://inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) puede resultar útil.


```python
intercept = ...
intercept
```

```python
grader.check("q2_2")
```

## 3. Investigando la línea de regresión
La pendiente y la intersección te dicen exactamente cómo se ve la línea de regresión. Para predecir el tiempo de espera para una erupción, multiplique la duración de la erupción por "pendiente" y luego agregue "intersección".

**Pregunta 3.1.** Calcule el tiempo de espera previsto para una erupción que dura 2 minutos y para una erupción que dura 5 minutos.


```python
two_minute_predicted_waiting_time = ...
five_minute_predicted_waiting_time = ...

# Here is a helper function to print out your predictions.
# Don't modify the code below.
def print_prediction(duration, predicted_waiting_time):
    print("After an eruption lasting", duration,
          "minutes, we predict you'll wait", predicted_waiting_time,
          "minutes until the next eruption.")

print_prediction(2, two_minute_predicted_waiting_time)
print_prediction(5, five_minute_predicted_waiting_time)
```

```python
grader.check("q3_1")
```

La siguiente celda traza la línea que va entre esos dos puntos, que es (un segmento de) la línea de regresión.

```python
plot_data_and_line(faithful, "duration", "wait",
                   [2, two_minute_predicted_waiting_time],
                   [5, five_minute_predicted_waiting_time])
```

**Pregunta 3.2.** Haga predicciones para el tiempo de espera después de cada erupción en la tabla `fiel`. (¡Por supuesto, sabemos exactamente cuáles fueron los tiempos de espera! Estamos haciendo esto para poder ver cuán precisas son nuestras predicciones). Coloque estos números en una columna en una nueva tabla llamada "predicciones_fieles". Su primera fila debería verse así:

|duración|espera|espera prevista|
|-|-|-|
|3.6|79|72.1011|

*Pista:* Tu respuesta puede ser solo una línea, aunque no estás limitado a una sola línea. No es necesario un bucle "for"; use aritmética de matrices en su lugar.

```python
faithful_predictions = ...
faithful_predictions
```

```python
grader.check("q3_2")
```

**Pregunta 3.3.** ¿Qué tan cerca estábamos? Calcule el *residual* de cada erupción en el conjunto de datos. El residual es el tiempo de espera real menos el tiempo de espera previsto. Agregue los residuos a `faithful_predictions` como una nueva columna llamada `residual` y nombre la tabla resultante `faithful_residuals`.

Como se menciona en el [Capítulo 15.6](https://inferentialthinking.com/chapters/15/6/Numerical_Diagnostics.html?highlight=residuals#average-of-residuals), una propiedad útil de los residuos producidos por la regresión lineal con una intersección es que la suma de los residuos siempre será cero. Siempre puedes verificar esto para asegurarte de haber calculado tus residuos correctamente. Tenga en cuenta que puede estar ligeramente alejado de cero debido a errores de redondeo.

¡Siéntete libre de copiar y pegar esta línea de código para comprobar tu trabajo! Tenga en cuenta que la respuesta puede ser *muy ligeramente* diferente de cero debido a errores de redondeo en Python.

```pitón
suma(fiel_residuals.columna('residual'))
```

*Pista:* Nuevamente, su código será mucho más simple si no usa un bucle `for`.


```python
residuals = ...
faithful_residuals = ...
faithful_residuals
```

```python
grader.check("q3_3")
```

Aquí hay un gráfico de los residuos que calculó. Cada punto corresponde a una erupción. Muestra hasta qué punto nuestra predicción sobreestimó o subestimó el tiempo de espera.

```python
faithful_residuals.scatter("duration", "residual", color="r")
```

Realmente no hay un patrón en los residuos, lo que confirma que era razonable intentar la regresión lineal. Es cierto que hay dos nubes separadas; la duración de las explosiones pareció dividirse en dos grupos distintos. Pero eso es sólo un patrón en la duración de las erupciones, no un patrón en la relación entre las duraciones de las erupciones y los tiempos de espera.

## 4. ¿Qué tan precisas son las diferentes predicciones?
Anteriormente, deberías haber descubierto que la correlación es bastante cercana a 1, por lo que la línea se ajusta bastante bien a los datos de entrenamiento. Eso significa que los residuos son en general pequeños (cerca de 0) en comparación con los tiempos de espera.

Podemos verlo visualmente trazando juntos los tiempos de espera y los residuos:

```python
# Just run this cell.
faithful_residuals.scatter("duration", "wait", label="actual waiting time", color="blue")
plots.scatter(faithful_residuals.column("duration"), faithful_residuals.column("residual"), label="residual", color="r")
plots.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time], label="regression line")
plots.legend(bbox_to_anchor=(1.7,.8));
```

Sin embargo, a menos que tenga una razón sólida para creer que el modelo de regresión lineal es verdadero, debe tener cuidado al aplicar su modelo de predicción a datos que sean muy diferentes de los datos de entrenamiento.

**Pregunta 4.1.** En "fieles", ninguna erupción duró exactamente 0, 2,5 o 60 minutos. Usando esta línea, ¿cuál es el tiempo de espera previsto para una erupción que dura 0 minutos? ¿2,5 minutos? ¿Una hora?

```python
zero_minute_predicted_waiting_time = ...
two_point_five_minute_predicted_waiting_time = ...
hour_predicted_waiting_time = ...

print_prediction(0, zero_minute_predicted_waiting_time)
print_prediction(2.5, two_point_five_minute_predicted_waiting_time)
print_prediction(60, hour_predicted_waiting_time)
```

```python
grader.check("q4_1")
```

**Pregunta 4.2.** Para cada predicción, indica si crees que es confiable y explica tu razonamiento.

_Escriba su respuesta aquí, reemplazando este texto._

## 5. Divide y vencerás

El siguiente gráfico de dispersión muestra la relación entre la duración de las erupciones en el eje x y el tiempo de espera hasta la siguiente erupción en el eje y.

Del diagrama de dispersión se desprende que hay dos grupos de puntos: uno con duraciones de alrededor de 2 y otro con duraciones de entre 3,5 y 5. Una línea vertical en 3 divide estos dos grupos.

```python
faithful.scatter("duration", "wait", label="actual waiting time", color="blue")
plots.plot([3, 3], [40, 100]);
```

La función "estandarizar" de la conferencia aparece a continuación, que toma una tabla con columnas numéricas y devuelve la misma tabla con cada columna convertida en unidades estándar.

```python
# Just run this cell.

def standard_units(any_numbers):
    "Convert any array of numbers to standard units."
    return (any_numbers - np.mean(any_numbers)) / np.std(any_numbers)

def standardize(table):
    """Return a table in which all columns of t are converted to standard units."""
    t_su = Table()
    for label in table.labels:
        t_su = t_su.with_column(label + ' (su)', standard_units(table.column(label)))
    return t_su
```

**Pregunta 5.1.** Calcule por separado el coeficiente de correlación *r* para todos los puntos con una duración inferior a 3 **y luego** para todos los puntos con una duración superior a 3. Para hacerlo, cree una función que calcule ` r` de una tabla, y luego pasarle dos tablas de puntos diferentes, llamadas `below_3` y `above_3`.

*Pista:* Puedes asumir que la tabla no tiene ningún valor de duración que sea exactamente 3.

```python
def corr_coeff(table):
    """Return the regression coefficient for columns 0 & 1."""
    t_su = standardize(table)
    ...

below_3 = ...
above_3 = ...
below_3_r = corr_coeff(below_3)
above_3_r = corr_coeff(above_3)
print("For points below 3, r is", below_3_r, "; for points above 3, r is", above_3_r)
```

```python
grader.check("q5_1")
```

**Pregunta 5.2.** Complete las funciones `slope_of` e `intercept_of` a continuación.

Cuando haya terminado, las funciones `wait_below_3` y `wait_above_3` deberían usar cada una una línea de regresión diferente para predecir un tiempo de espera durante una duración. La primera función debe usar la línea de regresión para todos los puntos con una duración inferior a 3. La segunda función debe usar la línea de regresión para todos los puntos con una duración superior a 3.

```python
def slope_of(table, r):
    """Return the slope of the regression line for table in original units.

    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    """
    ...

def intercept_of(table, r):
    """Return the intercept of the regression line for table in original units.

    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    """
    slope = slope_of(table, r)
    ...

below_3_slope = slope_of(below_3, below_3_r)
below_3_intercept = intercept_of(below_3, below_3_r)
above_3_slope = slope_of(above_3, above_3_r)
above_3_intercept = intercept_of(above_3, above_3_r)

def wait_below_3(duration):
    return below_3_slope * duration + below_3_intercept

def wait_above_3(duration):
    return above_3_slope * duration + above_3_intercept
```

```python
grader.check("q5_2")
```

El siguiente gráfico muestra las dos líneas de regresión diferentes, una para cada grupo, ¡junto con la línea de regresión original!

```python
faithful.scatter(0, 1)
plots.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time])
plots.plot([1, 3], [wait_below_3(1), wait_below_3(3)])
plots.plot([3, 6], [wait_above_3(3), wait_above_3(6)]);
```

**Pregunta 5.3.** Escriba una función `predict_wait` que tome una `duración` y devuelva el tiempo de espera previsto utilizando la línea de regresión adecuada, dependiendo de si la duración es inferior a 3 o superior a 3.

*Pista:* Considere la posibilidad de utilizar declaraciones condicionales.


```python
def predict_wait(duration):
    ...
```

```python
grader.check("q5_3")
```

Los tiempos de espera previstos para cada punto aparecen a continuación.

```python
faithful_pred_split = faithful.with_column('predicted', faithful.apply(predict_wait, 'duration'))
faithful_pred_split.scatter(0)
```

**Pregunta 5.4.** ¿Crees que las predicciones producidas por `predict_wait` serían más o menos precisas que las predicciones de la línea de regresión que creaste en la sección 2? ¿Cómo podrías saberlo?

**SOLUCIÓN:** Más preciso, porque cada línea es específica de los valores de su grupo. Para verificarlo, podríamos medir la magnitud promedio de los valores residuales.

La siguiente celda trazará los residuos de cada erupción en el conjunto de datos cuando tengamos una línea de regresión y dos líneas de regresión. También vemos la magnitud promedio de los valores residuales.

```python
# Just run this cell
faithful_pred_split_residuals = faithful_pred_split.with_column('residual', faithful_pred_split.column(1) - faithful_pred_split.column(2))
plots.scatter(faithful_residuals.column('duration'), faithful_residuals.column('residual'), label='one regression line')
plots.scatter(faithful_pred_split_residuals.column('duration'), faithful_pred_split_residuals.column('residual'), label='two regression lines');
plots.axis([1, 6, -15, 15])
plots.legend(bbox_to_anchor=(1.5,.8));
print("Average Magnitude of Residual Values for One Regression Line: ", np.mean(abs(faithful_residuals.column('residual'))))
print("Average Magnitude of Residual Values for Two Regression Lines: ", np.mean(abs(faithful_pred_split_residuals.column('residual'))))
```

La gráfica residual para los tiempos de espera cuando son predichos por dos líneas de regresión (roja) realmente no tiene un patrón, lo que confirma que también era apropiado usar la regresión lineal en nuestro escenario "Divide y vencerás". ¿Cómo se comparan las dos gráficas residuales?

-----
## ¡Hecho!

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab09/milo.png?raw=1" alt="Picture of a cordi" width="300"/>

**¡Milo** tuvo un gran Halloween!


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**


## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_0 (points: None)
```
>>> len(least_squares_order) == 4
True
```
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(least_squares_order).astype(int)) # Your ordering is incorrect.
'cd29f2d730e11535cef30ecc78640daa'
```


Tests for q1_3 (points: None)
```
>>> abs(sum(faithful_standard.column(0))) <= 1e-8
True
```
```
>>> round(duration_std, 2) == 1.14
True
```
```
>>> round(wait_std, 2) == 13.57 
True
```


Tests for q1_5 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(correlation)) # Your choice is incorrect.
'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q1_6 (points: None)
```
>>> -1 <= r <= 1
True
```
```
>>> round(r,3) == 0.901
True
```


Tests for q2_1 (points: None)
```
>>> (slope*13 - 100)/98 <= 0.5
True
```
```
>>> np.round(slope, 4) == 10.7296 # Make sure you are plugging in the SD_x and SD_y in the correct spots!
True
```


Tests for q2_2 (points: None)
```
>>> 33 < intercept < 34
True
```
```
>>> np.round(intercept, 3) == 33.474
True
```


Tests for q3_1 (points: None)
```
>>> round(two_minute_predicted_waiting_time, 3) == 54.934
True
```
```
>>> round(five_minute_predicted_waiting_time, 3) == 87.123
True
```


Tests for q3_2 (points: None)
```
>>> set(faithful_predictions.labels) == set(['duration', 'wait', 'predicted wait']) # Make sure your column labels are correct.
True
```
```
>>> abs(1 - np.mean(faithful_predictions.column(2))/100) <= 0.35
True
```


Tests for q3_3 (points: None)
```
>>> set(faithful_residuals.labels) == set(['duration', 'wait', 'predicted wait', 'residual']) # Make sure your column labels are correct.
True
```
```
>>> abs(sum(faithful_residuals.column('residual'))) <= 1e-8
True
```


Tests for q4_1 (points: None)
```
>>> 12 - zero_minute_predicted_waiting_time*1.4/4 <= 0.35
True
```
```
>>> zero_minute_predicted_waiting_time - intercept <= 1e5
True
```
```
>>> 2 - two_point_five_minute_predicted_waiting_time/35 <= 0.4
True
```
```
>>> (26 - hour_predicted_waiting_time/30)/10 <= 0.43
True
```


Tests for q5_1 (points: None)
```
>>> [below_3.num_rows, above_3.num_rows]
[97, 175]
```
```
>>> np.allclose([below_3_r, above_3_r], [0.290189526493, 0.372782225571])
True
```


Tests for q5_2 (points: None)
```
>>> np.allclose([wait_below_3(1), wait_below_3(3), wait_above_3(3), wait_above_3(6)], 
...             [47.902151605742517, 60.603197182023813, 72.965413990538366, 89.281859197449506])
True
```


Tests for q5_3 (points: None)
```
>>> np.allclose([predict_wait(x) for x in [1.5, 2.5]], 
...             [51.077412999812836, 57.427935787953487])
True
```
```
>>> np.allclose([predict_wait(x) for x in [3.5, 4.5]], 
...             [75.684821525023551, 81.123636593993922])
True
```
